# Sprint Challenge
## *Data Science Unit 4 Sprint 1*

After a week of Natural Language Processing, you've learned some cool new stuff: how to process text, how turn text into vectors, and how to model topics from documents. Apply your newly acquired skills to one of the most famous NLP datasets out there: [Yelp](https://www.yelp.com/dataset). As part of the job selection process, some of my friends have been asked to create analysis of this dataset, so I want to empower you to have a head start.  

The real dataset is massive (almost 8 gigs uncompressed). I've sampled the data for you to something more manageable for the Sprint Challenge. You can analyze the full dataset as a stretch goal or after the sprint challenge. As you work on the challenge, I suggest adding notes about your findings and things you want to analyze in the future.

## Challenge Objectives
Successfully complete all these objectives to earn full credit. 

**Successful completion is defined as passing all the unit tests in each objective.**  

Each unit test that you pass is 1 point. 

There are 5 total possible points in this sprint challenge. 


There are more details on each objective further down in the notebook.*
* <a href="#p1">Part 1</a>: Write a function to tokenize the yelp reviews
* <a href="#p2">Part 2</a>: Create a vector representation of those tokens
* <a href="#p3">Part 3</a>: Use your tokens in a classification model on yelp rating
* <a href="#p4">Part 4</a>: Estimate & Interpret a topic model of the Yelp reviews

____

# Before you submit your notebook you must first

1) Restart your notebook's Kernel

2) Run all cells sequentially, from top to bottom, so that cell numbers are sequential numbers (i.e. 1,2,3,4,5...)
- Easiest way to do this is to click on the **Cell** tab at the top of your notebook and select **Run All** from the drop down menu. 

3) Comment out the cell that generates a pyLDAvis visual in objective 4 (see instructions in that section). 
____



### Import Data

<!-- grade_id: cell-395851cd95d17235, points: 0, grade: false, solution: true -->
```python
import pandas as pd

# Load reviews from URL
data_url = 'https://raw.githubusercontent.com/LambdaSchool/data-science-practice-datasets/main/unit_4/unit1_nlp/review_sample.json'

# Import data into a DataFrame named df
# YOUR CODE HERE

df = pd.read_json(data_url, lines=True)

#raise NotImplementedError()
```

<!-- grade_id: cell-cb5006475e42b8f9, points: 1, grade: true, solution: false -->
```python
# Visible Testing
assert isinstance(df, pd.DataFrame), 'df is not a DataFrame. Did you import the data into df?'
assert df.shape[0] == 10000, 'DataFrame df has the wrong number of rows.'
```

```python
df.head()
```

Output:
```
business_id  cool                date  funny  \
0  nDuEqIyRc8YKS1q1fX0CZg     1 2015-03-31 16:50:30      0   
1  eMYeEapscbKNqUDCx705hg     0 2015-12-16 05:31:03      0   
2  6Q7-wkCPc1KF75jZLOTcMw     1 2010-06-20 19:14:48      1   
3  k3zrItO4l9hwfLRwHBDc9w     3 2010-07-13 00:33:45      4   
4  6hpfRwGlOzbNv7k5eP9rsQ     1 2018-06-30 02:30:01      0   

                review_id  stars  \
0  eZs2tpEJtXPwawvHnHZIgQ      1   
1  DoQDWJsNbU0KL1O29l_Xug      4   
2  DDOdGU7zh56yQHmUnL1idQ      3   
3  LfTMUWnfGFMOfOIyJcwLVA      1   
4  zJSUdI7bJ8PNJAg4lnl_Gg      4   

                                                text  useful  \
0  BEWARE!!! FAKE, FAKE, FAKE....We also own a sm...      10   
1  Came here for lunch Togo. Service was quick. S...       0   
2  I've been to Vegas dozens of times and had nev...       2   
3  We went here on a night where they closed off ...       5   
4  3.5 to 4 stars\n\nNot bad for the price, $12.9...       5   

                  user_id  
0  n1LM36qN
```

## Part 1: Tokenize Function
<a id="#p1"></a>

Complete the function `tokenize`. Your function should
- accept one document at a time
- return a list of tokens

You are free to use any method you have learned this week.

```python
# Optional: Consider using spaCy in your function. The spaCy library can be imported by running this cell.
# A pre-trained model (en_core_web_sm) has been made available to you in the CodeGrade container.
# If you DON'T need use the en_core_web_sm model, you can comment it out below.
import spacy
nlp = spacy.load('en_core_web_md')
```

Output:
```
/anaconda2/envs/U4-S1-NLP/lib/python3.7/site-packages/catalogue.py:138: DeprecationWarning: SelectableGroups dict interface is deprecated. Use select.
  for entry_point in AVAILABLE_ENTRY_POINTS.get(self.entry_point_namespace, []):
/anaconda2/envs/U4-S1-NLP/lib/python3.7/site-packages/catalogue.py:138: DeprecationWarning: SelectableGroups dict interface is deprecated. Use select.
  for entry_point in AVAILABLE_ENTRY_POINTS.get(self.entry_point_namespace, []):
/anaconda2/envs/U4-S1-NLP/lib/python3.7/site-packages/catalogue.py:126: DeprecationWarning: SelectableGroups dict interface is deprecated. Use select.
  for entry_point in AVAILABLE_ENTRY_POINTS.get(self.entry_point_namespace, []):
/anaconda2/envs/U4-S1-NLP/lib/python3.7/site-packages/catalogue.py:138: DeprecationWarning: SelectableGroups dict interface is deprecated. Use select.
  for entry_point in AVAILABLE_ENTRY_POINTS.get(self.entry_point_namespace, []):
```

```python
df.text[300]
```

Output:
```
"Went here at 5:30 recently at the suggestion of a downtown hipster friend. Place was small and packed, but had great selection of wines and beers that made the people I was with happy. Personally I'm a cocktail guy, so the beer and wine only deal is always a letdown for me. Service was passable, the place was small but well finished and is a small house divided into a couple of rooms. This is a place trhat you should be prepared to stand at the enitire time, girls with us in heels felt free to bith about that. The front room seats 4 on opposite sides on the room, the ante room seated about three and the back room had a table for eight, I think. That's it, unless you are lucky enough to grab one of the six or eight precious stools at the bar. Interesting mix of professionals and record store employees (pierced and tattaed and etc. and some reasonable eye candy from both camps) who were probably 4 hours early to see some shoe gazers play an off-tune set at Modified. The place has no sig
```

```python
import re

def only_alphanum(text):
    alphanum = re.sub('[^a-zA-Z 0-9]', '', text)
    alphanum = re.sub(r"\s+", " ", alphanum)
    return alphanum

print(only_alphanum(df.text[300]).split())
```

Output:
```
['Went', 'here', 'at', '530', 'recently', 'at', 'the', 'suggestion', 'of', 'a', 'downtown', 'hipster', 'friend', 'Place', 'was', 'small', 'and', 'packed', 'but', 'had', 'great', 'selection', 'of', 'wines', 'and', 'beers', 'that', 'made', 'the', 'people', 'I', 'was', 'with', 'happy', 'Personally', 'Im', 'a', 'cocktail', 'guy', 'so', 'the', 'beer', 'and', 'wine', 'only', 'deal', 'is', 'always', 'a', 'letdown', 'for', 'me', 'Service', 'was', 'passable', 'the', 'place', 'was', 'small', 'but', 'well', 'finished', 'and', 'is', 'a', 'small', 'house', 'divided', 'into', 'a', 'couple', 'of', 'rooms', 'This', 'is', 'a', 'place', 'trhat', 'you', 'should', 'be', 'prepared', 'to', 'stand', 'at', 'the', 'enitire', 'time', 'girls', 'with', 'us', 'in', 'heels', 'felt', 'free', 'to', 'bith', 'about', 'that', 'The', 'front', 'room', 'seats', '4', 'on', 'opposite', 'sides', 'on', 'the', 'room', 'the', 'ante', 'room', 'seated', 'about', 'three', 'and', 'the', 'back', 'room', 'had', 'a', 'table', 'for', 'e
```

<!-- grade_id: cell-3d570d5a1cd6cb64, points: 0, grade: false, solution: true -->
```python
def tokenize(document):
    doc = only_alphanum(document) #using our function
    doc = nlp(document)
    doc_tokens = []
    for token in doc:
        if (not token.is_punct) & (not token.is_stop) & (token.text != " "):
            doc_tokens.append(token.lemma_.strip().lower())
    return doc_tokens
#raise NotImplementedError()
```

<!-- grade_id: cell-02da164f6fbe730a, points: 1, grade: true, solution: false -->
```python
'''Testing'''
assert isinstance(tokenize(df.sample(n=1)["text"].iloc[0]), list), "Make sure your tokenizer function accepts a single document and returns a list of tokens!"
```

Output:
```
/anaconda2/envs/U4-S1-NLP/lib/python3.7/site-packages/catalogue.py:138: DeprecationWarning: SelectableGroups dict interface is deprecated. Use select.
  for entry_point in AVAILABLE_ENTRY_POINTS.get(self.entry_point_namespace, []):
/anaconda2/envs/U4-S1-NLP/lib/python3.7/site-packages/catalogue.py:138: DeprecationWarning: SelectableGroups dict interface is deprecated. Use select.
  for entry_point in AVAILABLE_ENTRY_POINTS.get(self.entry_point_namespace, []):
```

## Part 2: Vector Representation
<a id="#p2"></a>
1. Create a vector representation of the reviews (i.e. create a doc-term matrix).
2. Write a fake review and query for the 10 most similar reviews, print the text of the reviews. Do you notice any patterns?
    - Given the size of the dataset, use `NearestNeighbors` model for this. 

```python
corpus = list(df['text'].copy()) #working on a copy of the column with the docs and making it a list 
```

```python

```

<!-- grade_id: cell-0e96491cb529202c, points: 0, grade: false, solution: true -->
```python
%%time
# Create a vector representation of the reviews 
# Name that doc-term matrix "dtm"

# YOUR CODE HERE

from sklearn.feature_extraction.text import TfidfVectorizer
from sklearn.feature_extraction.text import CountVectorizer

# YOUR CODE HERE

def generate_dtm_CVectorizer(corpus):
    vect = CountVectorizer()
    dtm = vect.fit_transform(corpus)
    dtm = pd.DataFrame(dtm.todense(), columns=vect.get_feature_names())
    return dtm

def generate_dtm_tfidf(corpus):
    tfidf = TfidfVectorizer(stop_words='english', max_features=5000, 
                        ngram_range=(1,2), max_df=.98, min_df=.01, tokenizer=tokenize) #using the doc tokenizer we built

    dtm = tfidf.fit_transform(corpus) 

    dtm = pd.DataFrame(dtm.todense(), columns=tfidf.get_feature_names())
    
    return dtm

def get_nearest(dtm, dtm_index, my_nn):
    dtm_doc = [dtm.iloc[dtm_index].values]
    neigh_dist, neigh_index = my_nn.kneighbors(dtm_doc)
    df = pd.DataFrame(list(zip(neigh_index[0], neigh_dist[0])), columns=['index', 'cos_sim'])
    return df
    
def nearest_from_text(generator, corpus , text):
    my_corpus = corpus.copy()
    my_corpus.append(text)
    dtm = generator(my_corpus)
    nn = NearestNeighbors(n_neighbors=10, algorithm='kd_tree') #we are aiming at 10 nn so we set param at 11
    nn.fit(dtm)
    df = get_nearest(dtm, -1, nn)
    return df

dtm = generate_dtm_CVectorizer(corpus)
dtm.head()

### END SOLUTION


#raise NotImplementedError()
```

Output:
```
CPU times: user 1.47 s, sys: 906 ms, total: 2.37 s
Wall time: 2.73 s
   00  000  001695  007  00a  00am  00ish  00pm  01  011802  ...  誰も乗車しなくても  \
0   0    0       0    0    0     0      0     0   0       0  ...          0   
1   0    0       0    0    0     0      0     0   0       0  ...          0   
2   0    0       0    0    0     0      0     0   0       0  ...          0   
3   0    0       0    0    0     0      0     0   0       0  ...          0   
4   0    0       0    0    0     0      0     0   0       0  ...          0   

   豆腐花  質問にも丁寧に答えてくれましたし  車好きさんには  這是一個不錯的選擇  運転しない  運転中も英語で指導があります  食べ物はうまい  \
0    0                 0        0          0      0               0        0   
1    0                 0        0          0      0               0        0   
2    0                 0        0          0      0               0        0   
3    0                 0        0          0      0               0        0   
4    0                 0        0          0      0     
```

```python
dtm.columns[5000]
```

<!-- grade_id: cell-3d5bc610a8ec6b24, points: 0, grade: false, solution: true -->
```python
# Create and fit a NearestNeighbors model named "nn"
from sklearn.neighbors import NearestNeighbors

nn = NearestNeighbors(n_neighbors=10, algorithm='kd_tree') #we are aiming at 10 nn so we set param at 11
nn.fit(dtm)

#raise NotImplementedError()
```

Output:
```
NearestNeighbors(algorithm='kd_tree', leaf_size=30, metric='minkowski',
                 metric_params=None, n_jobs=None, n_neighbors=10, p=2,
                 radius=1.0)
```

<!-- grade_id: cell-c43704dcff67e99b, points: 1, grade: true, solution: false -->
```python
'''Testing.'''
assert nn.__module__ == 'sklearn.neighbors._unsupervised', ' nn is not a NearestNeighbors instance.'
assert nn.n_neighbors == 10, 'nn has the wrong value for n_neighbors'
```

<!-- grade_id: cell-496203e8746296ca, points: 0, grade: false, solution: true -->
```python
# Create a fake review and find the 10 most similar reviews

# YOUR CODE HERE

my_text = "It was the best coffee though super expensive for a residential area, severs were excellent and decorations are on spot"

print(nearest_from_text(generate_dtm_CVectorizer, corpus, my_text))

#raise NotImplementedError()
```

Output:
```
index   cos_sim
0  10000  0.000000
1   9898  4.472136
2   5129  4.472136
3   3543  4.472136
4   5087  4.582576
5   7352  4.690416
6   7553  4.690416
7   1356  4.690416
8   3405  4.795832
9   6810  4.795832
```

```python
print(get_nearest(dtm, -1, nn)) #chacking that we are not overwriting the original values
```

Output:
```
index    cos_sim
0   9999   0.000000
1   4355  24.758837
2    586  25.000000
3   6262  25.396850
4   8603  25.436195
5   3991  25.514702
6   7276  25.632011
7   4104  25.690465
8   5442  25.748786
9      2  25.806976
```

## Part 3: Classification
<a id="#p3"></a>
Your goal in this section will be to predict `stars` from the review dataset. 

1. Create a pipeline object with a sklearn `CountVectorizer` or `TfidfVector` and any sklearn classifier.
    - Use that pipeline to train a model to predict the `stars` feature (i.e. the labels). 
    - Use that Pipeline to predict a star rating for your fake review from Part 2. 



2. Create a parameter dict including `one parameter for the vectorizer` and `one parameter for the model`. 
    - Include 2 possible values for each parameter
    - **Use `n_jobs` = 1** 
    - Due to limited computational resources on CodeGrader `DO NOT INCLUDE ADDITIONAL PARAMETERS OR VALUES PLEASE.`
    
    
3. Train the entire pipeline with a GridSearch
    - Name your GridSearch object as `gs`

<!-- grade_id: cell-e2beb0252d274bba, points: 0, grade: false, solution: true -->
```python
from sklearn.pipeline import Pipeline
from sklearn.feature_extraction.text import TfidfVectorizer
from sklearn.neighbors import KNeighborsClassifier
from sklearn.model_selection import GridSearchCV

# Name the gridsearch instance "gs"

# YOUR CODE HERE

features = corpus #keep working with the copy "list format" of df['text']and not the original column as a preference
target = df['stars'].copy()

X = features
y = target

tfidf = TfidfVectorizer(stop_words='english', tokenizer=None) #no tokenizer func

knc = KNeighborsClassifier()

pipe = Pipeline([
    ("vect", tfidf),
    ("class", knc)
])

parameters = {
    'vect__max_df':(0.75, 0.95),
    #'vect__min_df':(0.01, 0.05),
    'class__n_neighbors':(3,5),
    #'class__leaf_size':(25,40)
}

gs = GridSearchCV(pipe, param_grid=parameters, n_jobs=1, cv=3, verbose=1)

gs.fit(X, y)




#raise NotImplementedError()
```

Output:
```
Fitting 3 folds for each of 4 candidates, totalling 12 fits
[Parallel(n_jobs=1)]: Using backend SequentialBackend with 1 concurrent workers.
[Parallel(n_jobs=1)]: Done  12 out of  12 | elapsed:   24.7s finished
GridSearchCV(cv=3, error_score=nan,
             estimator=Pipeline(memory=None,
                                steps=[('vect',
                                        TfidfVectorizer(analyzer='word',
                                                        binary=False,
                                                        decode_error='strict',
                                                        dtype=<class 'numpy.float64'>,
                                                        encoding='utf-8',
                                                        input='content',
                                                        lowercase=True,
                                                        max_df=1.0,
                                                        max_feat
```

```python
best_model = gs.best_estimator_
best_model
```

Output:
```
Pipeline(memory=None,
         steps=[('vect',
                 TfidfVectorizer(analyzer='word', binary=False,
                                 decode_error='strict',
                                 dtype=<class 'numpy.float64'>,
                                 encoding='utf-8', input='content',
                                 lowercase=True, max_df=0.75, max_features=None,
                                 min_df=1, ngram_range=(1, 1), norm='l2',
                                 preprocessor=None, smooth_idf=True,
                                 stop_words='english', strip_accents=None,
                                 sublinear_tf=False,
                                 token_pattern='(?u)\\b\\w\\w+\\b',
                                 tokenizer=None, use_idf=True,
                                 vocabulary=None)),
                ('class',
                 KNeighborsClassifier(algorithm='auto', leaf_size=30,
                                      metric='minkowski', metric_param
```

<!-- grade_id: cell-d07134c6fe5d056e, points: 1, grade: true, solution: false -->
```python
# Visible Testing
prediction = gs.predict(["I wish dogs knew how to speak English."])[0]
print (prediction)
prediction = gs.predict([my_text])[0] #predicting my random review
print (prediction)

assert prediction in df.stars.values, 'You gs object should be able to accept raw text within a list. Did you include a vectorizer in your pipeline?'
```

Output:
```
1
5
```

## Part 4: Topic Modeling

Let's find out what those yelp reviews are saying! :D

1. Estimate a LDA topic model of the review text
    - Set num_topics to `5`
    - Name your LDA model `lda`
2. Create 1-2 visualizations of the results
    - You can use the most important 3 words of a topic in relevant visualizations. Refer to yesterday's notebook to extract. 
3. In markdown, write 1-2 paragraphs of analysis on the results of your topic model

When you instantiate your LDA model, it should look like this: 

```python
lda = LdaModel(corpus=corpus,
               id2word=id2word,
               random_state=723812,
               num_topics = num_topics,
               passes=1
              )

```

__*Note*__: You can pass the DataFrame column of text reviews to gensim. You do not have to use a generator.

## Note about  pyLDAvis

**pyLDAvis** is the Topic modeling package that we used in class to visualize the topics that LDA generates for us.

You are welcomed to use pyLDAvis if you'd like for your visualization. However, **you MUST comment out the code that imports the package and the cell that generates the visualization before you submit your notebook to CodeGrade.** 

Although you should leave the print out of the visualization for graders to see (i.e. comment out the cell after you run it to create the viz). 

```python
from gensim import corpora
# Due to limited computationalresources on CodeGrader, use the non-multicore version of LDA 
from gensim.models.ldamodel import LdaModel
import gensim
import re
```

```python

import pyLDAvis
import pyLDAvis.gensim 
import matplotlib.pyplot as plt
%matplotlib inline 

```

### 1. Estimate a LDA topic model of the review tex

<!-- grade_id: cell-66331a185ff52f15, points: 0, grade: false, solution: true -->
```python
# Remember to read the LDA docs for more information on the various class attirbutes and methods available to you
# in the LDA model: https://radimrehurek.com/gensim/models/ldamodel.html

# don't change this value 
num_topics = 5

# use tokenize function you created earlier to create tokens 

tokens = df['text'].apply(tokenize) 

def filter_tokens(doc):
    doc_tokens = []
    for token in doc:
        if len(token) > 2:
            doc_tokens.append(token)
    return doc_tokens

tokens = tokens.apply(filter_tokens)

# create a id2word object (hint: use corpora.Dictionary)
id2word = corpora.Dictionary(tokens)
# create a corpus object (hint: id2word.doc2bow)
corpus = [id2word.doc2bow(text) for text in tokens]
# instantiate an lda model
lda = gensim.models.ldamulticore.LdaMulticore(corpus=corpus)
# YOUR CODE HERE
#raise NotImplementedError()
```

```python

```

#### Testing

<!-- grade_id: cell-5a3c181311134fa9, points: 1, grade: true, solution: false -->
```python
# Visible Testing
assert lda.get_topics().shape[0] == 5, 'Did your model complete its training? Did you set num_topics to 5?'
```

#### 2. Create 1-2 visualizations of the results

<!-- grade_id: cell-9b043e992fbd218c, points: 0, grade: false, solution: true -->
```python
import seaborn as sns
import matplotlib.pyplot as plt

# Use pyLDAvis (or a ploting tool of your choice) to visualize your results 

# YOUR CODE HERE
raise NotImplementedError()
```

#### 3. In markdown, write 1-2 paragraphs of analysis on the results of your topic model